IQCH regulates spermatogenesis by interacting with CaM to promote the expression of RNA-binding proteins

Author Response

The following is the authors’ response to the original reviews.

Reviewer #1 (Recommendations For The Authors):

1. Line 99-100 The authors claimed that IQCH is a novel IQ motif-containing protein, which is essential for spermiogenesis and fertilization. However, it is not clear if the currently published paper named an ancient testis-specific IQ motif containing H gene that regulates specific transcript isoform expression during spermatogenesis.

Response: Thanks to the reviewer’s comment. Yes, IQCH is the ancient testis-specific IQ motif containing H gene. According to the reviewer’s suggestion, we have revised the statement “Here, we revealed a testis-specific IQ motif containing H gene, IQCH, which is essential for spermiogenesis and fertilization” in Introduction part of revised manuscript.

1. Line 154-159 Immunofluorescence staining for the marker of the acrosome (peanut agglutinin: PNA) as well as the mitochondrial marker (Transcription Factor A, Mitochondrial: TFAM) was performed to confirm the deficiency of the acrosomes and mitochondria in the proband's spermatozoa. It seems that the spermatozoa acrosomes and mitochondria were severely defective in the proband. The authors should indicate IQCH's role in mitochondrial and acrosome function and IQCH's role in mitochondrial and acrosome function these points by explaining how IQCH is related to mitochondrial and acrosome deficiency. In addition to staining, other functional analyses should be performed to strengthen the claim of acrosome and mitochondrial defects.

Response: We appreciate the reviewer's valuable suggestion. Indeed, in our study, the results of multiomics analysis on WT and Iqch KO testes, including LC-MS/MS analysis, proteomic analysis, and RNA-seq analysis, found a potential role of IQCH in mitochondrial and acrosome function. GO analysis of these analysis indicated a significant enrichment in mitochondrial and acrosomal functions, including acrosomal vesicle, acrosome assembly, vesicle fusion with Golgi apparatus, mitochondrion organization, mitochondrial matrix, and so on. Among the enriched molecules, in particular, HNRNPK mainly expresses at Golgi phase and Cap phase (Biggiogera et al. 1993). ANXA7 is a calcium-dependent phospholipid-binding protein that is a negative regulator of mitochondrial apoptosis (Du et al. 2015). Loss of SLC25A4 results in mitochondrial energy metabolism defects in mice (Graham et al. 1997). Furthermore, we confirmed that IQCH interacted with HNRNPK, ANXA7, and SLC25A4 through Co-IP, and exhibited downregulation in the sperm of the Iqch KO mice by immunofluorescence and western blotting. Moreover, IQCH can bind to HNRPAB, which could influence the mRNAs level of Catsper-family, such as Catsper1, Catsper2, and Catsper3, which are crucial for acrosome development (Jin ZR et al). In addition, we also detected HNRPAB binding to Dnhd1, which affects mitochondria development (Tan C et al). Therefore, in addition to staining, the other functional analyses also have provided the evidence of acrosome and mitochondrial defects caused by IQCH absence.

1. Line 180-182 IQCH knockout mice were generated. It is not clear why Mut-IQCH mice were not generated to be consistent with the human sequencing data.

Response: Thanks for reviewer’s comments. To understand IQCH's impact on fecundity in mice, we employed CRISPR-Cas9 to generate mice encoding the orthologous variant of IQCH387+1_387+10del detected in humans. Regrettably, due to sequence complexity, the designed sgRNA's specificity and efficiency were low, hindering successful Iqch knock-in mouse construction. Considering IQCH387+1_387+10del results in absent expression, we pursued Iqch knockout mice to explore IQCH's role in spermatogenesis.

1. Line 241.Figure 5A Gene Ontology (GO) analysis of the IQCH-bound proteins revealed a particular enrichment in fertilization, sperm axoneme assembly, mitochondrial organization, calcium channel, and RNA processing. But these GO functions are not shown in Figure 5A. The entire Figure 5 should be revised to enhance readability.

Response: We sincerely apologize for the oversight. These GO functions were indeed identified during the analysis of IQCH-bound proteins. Regrettably, we unintentionally omitted these GO functions when creating the plots. We have revised the plots in Figure 5 in revised manuscript to enhance readability.

1. Line 242 "33 ribosomal proteins were identified (Fig. 5B), indicating that IQCH might be involved in protein synthesis". The authors should perform an analysis to support the claim of protein synthesis defects.

Response: Thanks to reviewer’s suggestions. Initially, we have supplemented Co-IP experiments to confirm the interaction between IQCH and three ribosomal proteins (RPL4, RPS3, and RPS7), chosen from a pool of 33 ribosomal proteins based on different protein scores (Figure R1). In addition, the proteomic analysis revealed 807 upregulated proteins and 1,186 downregulated proteins in KO mice compared to WT mice. We confirmed the key downregulated proteins by western blotting and immunofluorescence staining in the previous manuscript. These results indicated that IQCH might interact with ribosomal proteins to regulate protein expression. Naturally, the regulation of protein synthesis by IQCH requires further experiments for confirmation in future studies.

The interaction between IQCH and ribosomal proteins.

Co-IP assays confirmed that IQCH interacted with RPL4, RPS3, and RPS7 in WT mouse sperm.

1. Line 244 The authors mentioned too many GO functions without focus.

Response: Following reviewer’s suggestions, we have simplified IQCH-associated GO functions in the revised manuscript.

1. Figure 6, there are no negative controls in all co-IP experiments. Band sizes are not marked. Thus, all data can't be evaluated. This also raises concern about whether the LC-MS/MS experiment to identify IQCH interacting protein was well-controlled? All co-IP experiments were poorly designed to draw any conclusion.

Response: Thanks to reviewer’s comments. We have supplemented negative controls in all Co-IP experiments and provided band sizes in Figure 6 in revised manuscript.

1. The authors mentioned that IQCH can bind to CaM. But they didn't detect CaM protein in Figure 5. Did the LC-MS/MS experiment really work?

Response: Thanks to reviewer’s comments. We detected the interaction of CaM protein with IQCH in the LC-MS/MS experiment analysis, which has been submitted as new Data S1 in the revised manuscript. We also confirmed their binding in mouse sperm by Co-IP experiment and immunofluorescence staining, which results were shown in Figure 6 and Figure S10 in the previous study.

1. Figure 6D. Because IQCH is lost in Iqch KO sperm, what is the point of showing in the Co-IP assay that CaM does not bind to IQCH in Iqch KO sperm?

Response: Following reviewer’s suggestions, we have deleted the results of Co-IP assay that CaM could not bind to IQCH in Iqch KO sperm.

1. Figure 6E. The Co-IP assay does not support the authors' claim that the decreased expression of HNRPAB was due to the reduced binding of IQCH and CaM by the knockout of IQCH or CaM.

Response: Thanks to reviewer’s expert comments. Indeed, the results of Figure 6E confirmed the interaction of IQCH and CaM in K562 cells, and also showed that the expression of HNRPAB was reduced when IQCH or CaM was knocked down, suggesting that IQCH or CaM might regulate HNRPAB expression. While in Figure 6F, the downregulation of HNRPAB caused by knocking down IQCH (or CaM) cannot be rescued when overexpressed CaM (or IQCH), indicating that CaM (or IQCH) cannot mediate HNRPAB expression alone. Therefore, the reduced expression of HNRPAB in Figure 6E might result from the weakened interaction between IQCH and CaM, but not a superficial downregulation of IQCH or CaM expression. To avoid the confusion, we have modified the relevant description in the revied manuscript.

Reviewer #2 (Recommendations For The Authors):

Major comments:

1. Lines 117 and 129: Please provide the reference number (NM_xxx.x) for the IQCH isoform that was used to interpret this variant. This is key information. Also, please provide the predicted truncation consequence caused by this splicing variant to IQCH protein.

Response: Thanks to reviewer’s suggestions. We have added reference number (NM_0010317152) of IQCH in manuscript. We employed splice site prediction tools, such as SpliceAI, RDDC, and varSEAK, to assess the expression consequences of this IQCH splicing variant. These tools couldn't anticipate the outcome of this splicing variant. However, the results of minigene splicing assay showed that the IQCH c.387+1_387+10del resulted in degradation of IQCH.

1. Figure 1A: The deleted sequence indicated by the red box does not match IQCH c.387+1_387+10del. Please show a plot of the exon-intron boundary under the Sanger sequencing results of the WT allele.

Response: Thanks to reviewer’s suggestions. We are sorry for the use of non-standard descriptions about the results of Sanger sequencing. According to the HGVS nomenclature (Figure R2), we have modified the red box to match IQCH c.387+1_387+10del and have added the exon-intron boundary in Figure 1A accordingly.

HGVS nomenclature description of the IQCH variant.

The picture showed a detailed HGVS nomenclature description of IQCH c.387+1_387+10del.

Minor comments:

a) Manuscript title: It is suggested to change the title to "IQCH regulates spermatogenesis by interacting with CaM to promote the expression of RNA-binding proteins".

Response: According to reviewer’s suggestions, we have modified the title as “IQCH regulates spermatogenesis by interacting with CaM to promote the expression of RNA-binding proteins”.

b) Line 116: Please introduce the abbreviation WES. Also, please introduce the other abbreviations (such as WT, SEM, TEM, etc.) the first time they appear.

Response: Thanks to reviewer’s suggestions. We have provided the full explanations for all abbreviations upon their initial appearance.

c) Line 140, "Nonfunctional IQCH": Due to "the lack of IQCH expression" in Line 137, should "Nonfunctional IQCH" be changed into "IQCH deficiency"?

Response: Thanks for reviewer’s the detailed review. We have modified this title in Results part of the revised manuscript as followed: “IQCH deficiency leads to sperm with cracked axoneme structures accompanied by defects in the acrosome and mitochondria”

d) The information on the following references is incomplete: Sechi et al., Tian et al., Wang et al., and Xu et al. Please provide issue/page/article numbers.

Response: We are sorry for our oversight. We have provided the missing issue/page/article numbers for the references.

e) The title of Figure 1: Please emphasize that the male infertile-associated variant is "homozygous".

Response: Thanks to reviewer’s suggestions. We have revised the title of Figure 1 to emphasize the homozygous variant as follows: “Identification of a homozygous splicing mutation in IQCH in a consanguineous family with male infertility”.

f) Table 1: Please provide the reference paper for the normal values. Response: We appreciate the reviewer's detailed checks. We have provided the reference paper for the normal values in Table 1.

g) Figure 5F is distorted. Please make sure that it is a perfect circle.

Response: Thanks to reviewer’s suggestions. We have revised both the graphical representation and layout of Figure 5 in revised manuscript to make sure the readability.

Reviewer #3 (Recommendations For The Authors):

While the writing is generally clear, there are multiple examples of where the writing could be improved for clarity.

1. While some terms are defined throughout the manuscript, many abbreviations are not defined upon their first mention, such as WES, RT-PCR, TYH, HTF, KSOM, KEGG, RIPA, PMSE, SDS-PAGE, H&L, and HRP.

Response: Thanks to reviewer’s suggestions. We have provided the full explanations for all abbreviations upon their initial appearance.

1. On line 44, the claim that spermatogenesis is the "most complex biological process" is rather subjective and hard to support with concrete data.

Response: Thanks to reviewer’s suggestions. We have modified this description in the Introduction section as follow: “Spermatogenesis is one of the most complex biological process in male organisms and functions to produce mature spermatozoa from spermatogonia in three phases: (i) spermatocytogenesis (mitosis), (ii) meiosis, and (iii) spermiogenesis.”

1. On line 54, I think the authors meant "heterogeneous," not "heterologous."

Response: Thanks to reviewer’s comment. We have changed “heterologous” into “heterogeneous”.

1. On line 156, I think the authors meant "deficiency," not "deficient."

Response: Thanks to reviewer’s comment. We are sorry to make this mistake. We have made the correction in the revised version of the manuscript.

1. On line 300, K562 cells are mentioned, but neither in the Methods nor the Results are any details about the biological origin of these cells (or rationale for their use other than co-expression of IQCH and CaM) provided.

Response: Thanks to reviewer’s suggestion. K562 cell line is a human leukemia cell line and is enriched in the expression of IQCH and CaM, we thus opted to use this cell line for an easier knockdown of IQCH and CaM. We have supplemented the details about the biological origin of these cells in Method section of revised manuscript.

1. For the Results section describing Figure 6H, it would be nice to provide some explanation of the results of ICHQ overexpression alone relative to control situations and not just relative to the delta-IQ version or relative to simultaneous CaM manipulation.

Response: According to the reviewer’s suggestion, we have supplemented the co-transfection of control and CaM plasmids in HEK293T cells, and the results showed that the expression of HNRPAB in cells co-transfected with control and CaM plasmids was similar to that of co-transfected with IQCH (△IQ) /CaM plasmids, but was lower than that in the cells overexpressing the WT-IQCH and CaM plasmids, confirming the nonfunction of IQCH (△IQ) plasmids. We have shown the results in Figure 6H in the revised manuscript.

1. The sentence on lines 352-354 is confusing.

Response: We apologize for any confusion caused by the sentence in question. We have revisited the sentence and made appropriate revisions to enhance its clarity as follows: “Our findings suggest that the fertilization function is the main action of IQ motif-containing proteins, while each specific IQ motif-containing protein also has its own distinct role in spermatogenesis.”

1. The use of "employee" on line 371 is awkward and not very scientific.

Response: Thanks to reviewer’s comment. We have changed “employee” in to “downstream effector protein” on line 376

Changes from the earlier version to the revised version

Abstract

Abstract

IQ motif-containing proteins can be recognized by calmodulin (CaM) and are essential for many biological processes. However, the role of IQ motif-containing proteins in spermatogenesis is largely unknown. In this study, we identified a loss-of-function mutation in the novel gene IQ motif-containing H (IQCH) in a Chinese family with male infertility, characterized by a cracked flagellar axoneme and abnormal mitochondrial structure. To verify the function of IQCH, Iqch-knockout mice were generated by CRISPR-Cas9 technology which reproduced the human phenotypes. Mechanistically, IQCH can bind to CaM and then regulate the expression of RNA-binding proteins (especially HNRPAB), which are indispensable for spermatogenesis. Collectively, this study firstly unveiled the function of IQCH, expanded the role of IQ motif-containing proteins in reproductive processes, and provided important guidance for genetic counseling and gene diagnosis for male infertility.

Introduction

Earlier version: Spermatogenesis is the most complex biological process in male organisms and functions to produce mature spermatozoa from spermatogonia in three phases: (i) spermatocytogenesis (mitosis), (ii) meiosis, and (iii) spermiogenesis (Hess and Renato de Franca 2008). This delicate process can be easily disturbed and further cause male reproductive disorders. Male infertility affects 7% of men in the general population, and its causes vary, ranging from anatomical or genetic abnormalities, systemic diseases, infections, trauma, iatrogenic injury, and gonadotoxins (Krausz et al. 2018). Approximately 10% of the human genome is related to reproductive processes; thus, male infertility is often predicted to be largely genetic in origin, whereas only 4% of infertile men are diagnosed with a distinct genetic cause (Krausz and Riera-Escamilla 2018). Genetic causes are highly heterologous and involve chromosomal abnormalities, point mutations in single genes, copy number variations, microRNA dysfunction, and polygenic defects (Meschede and Horst 1997, Traven et al. 2017). The highest percentage of known genetic factors that account for up to 25% of male infertility are karyotype anomalies, Y chromosome microdeletions, and CFTR mutations, which are mostly associated with azoospermia (Krausz and Riera-Escamilla 2018). However, in a relatively high proportion of infertile men (40%), the etiology cannot be recognized and is also referred to as idiopathic. With the help of assisted reproductive techniques (ART), some men have the chance to reproduce, however, there is a risk of passing on undetermined genetic abnormalities. In addition, genetic defects leading to fertilization failure and embryo development arrest cannot be effectively rescued by ART; thus, discovering novel genetic factors and further confirming their molecular mechanisms are of clinical importance.

Revised version: Spermatogenesis is one of the most complex biological process in male organisms and functions to produce mature spermatozoa from spermatogonia in three phases: (i) spermatocytogenesis (mitosis), (ii) meiosis, and (iii) spermiogenesis (Hess and Renato de Franca 2008). This delicate process can be easily disturbed and further cause male reproductive disorders. Male infertility affects 7% of men in the general population, and its causes vary, ranging from anatomical or genetic abnormalities, systemic diseases, infections, trauma, iatrogenic injury, and gonadotoxins (Krausz et al. 2018). Approximately 10% of the human genome is related to reproductive processes; thus, male infertility is often predicted to be largely genetic in origin, whereas only 4% of infertile men are diagnosed with a distinct genetic cause (Krausz and Riera-Escamilla 2018). Genetic causes are highly heterogeneous and involve chromosomal abnormalities, point mutations in single genes, copy number variations, microRNA dysfunction, and polygenic defects (Meschede and Horst 1997, Traven et al. 2017). The highest percentage of known genetic factors that account for up to 25% of male infertility are karyotype anomalies, Y chromosome microdeletions, and CFTR mutations, which are mostly associated with azoospermia (Krausz and Riera-Escamilla 2018). However, in a relatively high proportion of infertile men (40%), the etiology cannot be recognized and is also referred to as idiopathic. With the help of assisted reproductive techniques (ART), some men have the chance to reproduce, however, there is a risk of passing on undetermined genetic abnormalities. In addition, genetic defects leading to fertilization failure and embryo development arrest cannot be effectively rescued by ART; thus, discovering novel genetic factors and further confirming their molecular mechanisms are of clinical importance.

Calcium (Ca) is an essential element that acts as a universal intracellular second messenger. Ca is indispensable for many physiological processes in male reproduction, including spermatogenesis, sperm movement, capacitation, hyperactivation, acrosome reaction, chemotaxis, and fertilization (Valsa et al. 2015). Therefore, Ca deficiency is mentioned to be a main contributor to male infertility (Beigi Harchegani et al. 2019). Calmodulin (CaM) is defined as a major Ca sensor that activates many kinds of enzymes in response to an increase in intracellular Ca2+ by interacting with a diverse group of cellular proteins (Klee et al. 1980, Means et al. 1982). CaM binds to proteins through recognition motifs, including the short CaM-binding motif containing conserved Ile and Gln residues (IQ motif) for Ca2+-independent binding and two related motifs for Ca2+-dependent binding (Nie et al. 2009). The identified IQ motif-containing proteins include a good range of proteins with biological functions, one of which is sperm surface proteins, suggesting that IQ motif-containing proteins might play a potential role in male reproductive processes (Wen et al. 1999, Dolmetsch et al. 2001). Moreover, it has been revealed that CaM is predominantly expressed in the mammalian testis, from the pachytene to meiotic division stages (Smoake et al. 1974, Kakiuchi et al. 1982, YAMAMOTO 1985, Sano et al. 1987, Kägi et al. 1988, Moriya et al. 1995). However, limited IQ motif-containing proteins have been reported to be responsible for male fertility. It has been shown that the IQ motif-containing D (IQCD) is primarily accumulated in the acrosome area of spermatids during mouse spermiogenesis, and the acrosome reaction was inhibited in human spermatozoa by anti-IQCD antibody, suggesting a potential function of IQCD in fertilization and the acrosome reaction (Zhang et al. 2019). Moreover, Iqcf1 knockout (KO) male mice showed significantly less fertility, which was related to reduced sperm motility and acrosome reaction (Fang et al. 2015). Importantly, Iqcg is required for mouse spermiogenesis, which is attributable to its role in flagellar formation and/or function (Harris et al. 2014). Humans and mice without IQCN presented failed fertilization rates related to manchette assembly defects (Dai et al. 2022). Although reliable findings provide substantial clues for IQ motif-containing proteins participating in male reproductive processes, the relevant regulatory mechanism and the function of many other IQ motif-containing proteins in spermatogenesis have not been determined thus far.

Earlier version: Here, we revealed a novel IQ motif-containing protein, IQCH, which is essential for spermiogenesis and fertilization. Disrupting IQCH leads to deficient acrosome action and abnormal structure of the axoneme and mitochondria in both humans and mice. Moreover, it is suggested that the interaction of IQCH and CaM is a prerequisite for IQCH function, which further regulates the expression of RNA-binding proteins (especially HNRPAB) during spermatogenesis. Collectively, our findings unveiled a novel genetic diagnostic indicator of male infertility, and the uncovered mechanism of IQCH in spermatogenesis might shed new light on the treatment of this disease.

Revised version: Here, we revealed a testis-specific IQ motif containing H gene, IQCH, which is essential for spermiogenesis and fertilization. Disrupting IQCH leads to deficient acrosome action and abnormal structure of the axoneme and mitochondria in both humans and mice. Moreover, it is suggested that the interaction of IQCH and CaM is a prerequisite for IQCH function, which further regulates the expression of RNA-binding proteins (especially HNRPAB) during spermatogenesis. Collectively, our findings unveiled a novel genetic diagnostic indicator of male infertility, and the uncovered mechanism of IQCH in spermatogenesis might shed new light on the treatment of this disease.

Results

Earlier version: Identification of a novel IQ motif-containing protein, IQCH, involved in male fertility

Revised version: Identification of a novel homozygous variant of IQCH involved in male infertility

Earlier version: A 33-year-old man from a consanguineous family with primary infertility for four years was recruited in our study (Fig. 1A). Semen analysis of this patient revealed an inordinately decreasing sperm motility and count, as well as abnormal sperm morphology (Table 1). In addition to the proband, his great uncle was also affected, he had never conceived during sexual intercourse without contraception in his marriage. We further explored the possible genetic cause through WES analysis, and a homozygous deletion variant in IQCH (c.387+1_387+10del) was identified in the proband. This variant is rare in the general human population according to the 1000 Genomes Project (0.007%), ExAC Browser (0.059%), and gnomAD databases (0.024%). We further clarified the putative contribution of this variant in this family by Sanger sequencing. Noticeably, this homozygous variant was affirmed in his infertile great-uncle, and the fertile parents of the proband carried the heterozygote of this variant (Fig. 1A).

Revised version: A 33-year-old man from a consanguineous family with primary infertility for four years was recruited in our study (Fig. 1A). Semen analysis of this patient revealed an inordinately decreasing sperm motility and count, as well as abnormal sperm morphology (Table 1). In addition to the proband, his great uncle was also affected, he had never conceived during sexual intercourse without contraception in his marriage. We further explored the possible genetic cause through whole-exome sequencing (WES) analysis, and a homozygous deletion variant in IQCH (c.387+1_387+10del, NM_0010317152) was identified in the proband. This variant is rare in the general human population according to the 1000 Genomes Project (0.007%), ExAC Browser (0.059%), and gnomAD databases (0.024%). We further clarified the putative contribution of this variant in this family by Sanger sequencing. Noticeably, this homozygous variant was affirmed in his infertile great-uncle, and the fertile parents of the proband carried the heterozygote of this variant (Fig. 1A).

Earlier version: Identification of a splicing mutation in IQCH in a consanguineous family with male infertility.

Revised version: Identification of a homozygous splicing mutation in IQCH in a consanguineous family with male infertility.

(A) Pedigree analysis of the consanguineous family with two infertile males (II-1 and IV-1), with the black arrow pointing to the proband. The affected males exhibited the homozygous variant in IQCH by Sanger sequencing. Sequence chromatograms are shown below the pedigree. (B) (i) The electrophoresis results of the minigene assay show a decrease in the molecular weight of the RT-PCR products generated from the Mut-IQCH (263 bp) compared with the WT-IQCH (381 bp). (ii) Sanger sequencing of the complementary DNA of the splicing mutation showing the deletion of exon 4 in IQCH. (iii) The pattern diagram demonstrating the splicing effects caused by the IQCH mutation. (iv) Western blotting results showed that the Mut–cIQCH plasmids did not express IQCH. NC, negative control. Three independent experiments were performed. (C) Immunofluorescence staining showed that the expression of IQCH was barely detected in the proband’s sperm compared with the control (blue, DAPI; green, α-Tubulin; red, IQCH; scale bars, 5 μm).

Earlier version: Moreover, we used the minigene splicing assay to examine the effect of this variant on IQCH mRNA splicing. The electrophoresis results revealed that WT-IQCH yielded one transcript with a size of 381[bp, whereas the Mut-IQCH resulted in one strong band of 263[bp (Fig. 1B, i). Sanger sequencing further showed that the variant resulted in the absence of the whole exon 4, which was expected to translate into a truncated protein (Fig. 1B, ii and iii). Therefore, we further constructed a plasmid containing the aberrant cDNA sequence of IQCH caused by the splicing mutation to verify the effect of the variant on protein expression. The western blotting results showed that the Mut-cIQCH plasmid did not express IQCH while the WT-cIQCH plasmid did (Fig. 1B, iv). We further conducted immunofluorescence staining of spermatozoa from the proband and the control. We barely detected the expression of IQCH protein in the spermatozoa from the proband compared to the control (Fig. 1C). Taken together, these results reveal that this identified IQCH variant is responsible for the splicing abnormality and further causes the lack of IQCH expression, which is likely the genetic cause of male infertility in this family.

Revised version: Moreover, we used the minigene splicing assay to examine the effect of this variant on IQCH mRNA splicing. The electrophoresis results revealed that wildtype-IQCH (WT-IQCH) yielded one transcript with a size of 381[bp, whereas the mutant-IQCH (Mut-IQCH) resulted in one strong band of 263[bp (Fig. 1B, i). Sanger sequencing further showed that the variant resulted in the absence of the whole exon 4, which was expected to translate into a truncated protein (Fig. 1B, ii and iii). Therefore, we further constructed a plasmid containing the aberrant cDNA sequence of IQCH caused by the splicing mutation to verify the effect of the variant on protein expression. The western blotting results showed that the Mut-cIQCH plasmid did not express IQCH while the WT-cIQCH plasmid did (Fig. 1B, iv). We further conducted immunofluorescence staining of spermatozoa from the proband and the control. We barely detected the expression of IQCH protein in the spermatozoa from the proband compared to the control (Fig. 1C). Taken together, these results reveal that this identified IQCH variant is responsible for the splicing abnormality and further causes the lack of IQCH expression, which is likely the genetic cause of male infertility in this family.

Earlier version: Nonfunctional IQCH leads to sperm with cracked axoneme structures accompanied by defects in the acrosome and mitochondria

Revised version: IQCH deficiency leads to sperm with cracked axoneme structures accompanied by defects in the acrosome and mitochondria

Earlier version: A comprehensive sperm morphology analysis was further conducted on the proband (the semen was not available from his great-uncle). Papanicolaou staining revealed that the spermatozoa displayed multiple flagellar morphological abnormalities, such as coiled, bent, irregular, and even fractured tails (Fig. 2A). These morphological anomalies were confirmed more precisely by SEM (Fig. 2B). Significantly, more subtle abnormalities, such as axoneme exposure, bending, and cracking, were further identified between the midpiece and the principal piece of the proband’s spermatozoa (Fig. 2B). As expected, a deranged or incomplete ‘9+2’ microtubule structure of the flagella were observed from the proband by TEM analysis (fig. S1, A). In addition, we detected ultrastructural defects in the spermatozoa nucleus, including irregular shape, large vacuoles, and deformed acrosomes (Fig. 2C). The mitochondria of the spermatozoa had an abnormal arrangement and enlarged diameter (Fig. 2C).

Revised version: A comprehensive sperm morphology analysis was further conducted on the proband (the semen was not available from his great-uncle). Papanicolaou staining revealed that the spermatozoa displayed multiple flagellar morphological abnormalities, such as coiled, bent, irregular, and even fractured tails (Fig. 2A). These morphological anomalies were confirmed more precisely by scanning electron microscopy (SEM) analysis (Fig. 2B). Significantly, more subtle abnormalities, such as axoneme exposure, bending, and cracking, were further identified between the midpiece and the principal piece of the proband’s spermatozoa (Fig. 2B). As expected, a deranged or incomplete ‘9+2’ microtubule structure of the flagella was observed from the proband by transmission electron microscopy (TEM) analysis (fig. S1, A). In addition, we detected ultrastructural defects in the spermatozoa nucleus, including irregular shape, large vacuoles, and deformed acrosomes (Fig. 2C). The mitochondria of the spermatozoa had an abnormal arrangement and enlarged diameter (Fig. 2C).

Abnormal flagellar morphology and defective acrosomes and mitochondria in the infertile patient.

(A and B) The Papanicolaou staining (A) and scanning electron microscopy (SEM) (B) results showed flagellar morphological abnormalities (scale bars in A, 5 μm; scale bars in B, 2.5 μm). The dotted boxes and red arrowheads denote the axoneme cracking and exposure. (C) Transmission electron microscopy (TEM) results showing the deformed acrosomes and abnormal arrangement and diameter mitochondria (scale bars, 500 nm). The red arrowheads denote the abnormal mitochondria. N, nucleus; M, mitochondria. (D and E) Defects of the acrosome and mitochondria were observed in the proband’s sperm by PNA (D) and TFAM (E) staining (blue, DAPI; green, α-Tubulin; red, PNA or TFAM; scale bars, 5 μm). PNA, peanut agglutinin; TFAM, Transcription Factor A, Mitochondrial.

Earlier version: Moreover, we performed immunofluorescence staining for the marker of the acrosome (peanut agglutinin: PNA) as well as the mitochondrial marker (Transcription Factor A, Mitochondrial: TFAM) to confirm the deficient of the acrosomes and mitochondria in the proband’s spermatozoa. The results of the PNA and TFAM staining suggested that the spermatozoa acrosomes and mitochondria were severely defective in the proband compared to the control (Fig. 2, D and E). In addition, we performed SEPTIN4 (SEPT4) staining, a functional marker of the annulus, to explore whether the flagellar fracture at the joint between the middle piece and principal piece was caused by a nonfunctional annulus. The results showed that there was no significant difference in the SEPT4 signal of the spermatozoa between the proband and the control (fig. S1, B).

Revised version: Moreover, we performed immunofluorescence staining for the marker of the acrosome (peanut agglutinin: PNA) as well as the mitochondrial marker (Transcription Factor A, Mitochondrial: TFAM) to confirm the deficiency of the acrosomes and mitochondria in the proband’s spermatozoa. The results of the PNA and TFAM staining suggested that the spermatozoa acrosomes and mitochondria were severely defective in the proband compared to the control (Fig. 2, D and E). In addition, we performed SEPTIN4 (SEPT4) staining, a functional marker of the annulus, to explore whether the flagellar fracture at the joint between the middle piece and principal piece was caused by a nonfunctional annulus. The results showed that there was no significant difference in the SEPT4 signal of the spermatozoa between the proband and the control (fig. S1, B).

Impairment of male fertility in mice without Iqch

Earlier version: To further consolidate the role of IQCH in male reproduction, we first explored the pattern of Iqch expression in mice by qPCR and found that Iqch was predominantly expressed in the mouse testis compared to other organs (fig. S2, A). Additionally, we investigated the temporal expression of Iqch in mouse testes on different postnatal days. The results revealed that the expression of Iqch showed a significant increase on postnatal Day 21, peaked on postnatal Day 35, and then presented stable expression (fig. S2, A). To better understand the role of IQCH during spermatogenesis, we performed immunofluorescence staining of the germ cells at different developmental stages in human and mouse testes. The exact localization of IQCH shared the same pattern in human and mouse spermatogenesis, being mainly detected in the cytoplasm of spermatocytes and round spermatids and the flagella of late spermatids (fig. S2, B and C). The testis-enriched expression of Iqch/IQCH suggested its potential role in spermatogenesis.

Revised version: To further consolidate the role of IQCH in male reproduction, we first explored the pattern of Iqch expression in mice by quantitative real-time PCR (qPCR) and found that Iqch was predominantly expressed in the mouse testis compared to other organs (fig. S2, A). Additionally, we investigated the temporal expression of Iqch in mouse testes on different postnatal days. The results revealed that the expression of Iqch showed a significant increase on postnatal Day 21, peaked on postnatal Day 35, and then presented stable expression (fig. S2, A). To better understand the role of IQCH during spermatogenesis, we performed immunofluorescence staining of the germ cells at different developmental stages in human and mouse testes. The exact localization of IQCH shared the same pattern in human and mouse spermatogenesis, being mainly detected in the cytoplasm of spermatocytes and round spermatids and the flagella of late spermatids (fig. S2, B and C). The testis-enriched expression of Iqch/IQCH suggested its potential role in spermatogenesis.

Earlier version: We next generated Iqch knockout (KO) mice using clustered regularly interspaced short palindromic repeat CRISPR-Cas9 technology. We used a guide RNA targeting exons 2 through 3 of Iqch to achieve the knockout (fig. S3, A). PCR, RT-PCR, and western blotting analysis further confirmed the success of the Iqch KO mouse construction (fig. S3, B). The Iqch KO mice showed no overt abnormalities in their development or behavior. The Iqch KO female mice showed normal fertility, and hematoxylin-eosin (HE) staining further showed that the Iqch KO female mice had normal follicle development compared to the WT mice (fig. S3, C-E). However, the fertility of the KO male mice was significantly reduced compared to that of the WT mice, including a reduction in the pregnancy rate and litter size (Fig. S3, F). There was no detectable difference in the testis/body weight ratio between the WT and Iqch KO mice (fig. S3, G), and the histology of testes and epididymides in the Iqch KO mice showed no obvious abnormalities compared to the WT mice (fig. S4, A and B).

Revised version: We next generated Iqch knockout (KO) mice using clustered regularly interspaced short palindromic repeat (CRISPR)-Cas9 technology. We used a guide RNA targeting exons 2 through 3 of Iqch to achieve the knockout (fig. S3, A). Polymerase chain reaction (PCR), reverse transcription-PCR (RT-PCR), and western blotting analysis further confirmed the success of the Iqch KO mouse construction (fig. S3, B). The Iqch KO mice showed no overt abnormalities in their development or behavior. The Iqch KO female mice showed normal fertility, and hematoxylin-eosin (H&E) staining further showed that the Iqch KO female mice had normal follicle development compared to the WT mice (fig. S3, C-E). However, the fertility of the KO male mice was significantly reduced compared to that of the WT mice, including a reduction in the pregnancy rate and litter size (Fig. S3, F). There was no detectable difference in the testis/body weight ratio between the WT and Iqch KO mice (fig. S3, G), and the histology of testes and epididymides in the Iqch KO mice showed no obvious abnormalities compared to the WT mice (fig. S4, A and B).

However, the results of the computer-assisted sperm analysis (CASA) affirmed that sperm motility was significantly reduced in the KO mice, and the sperm count was slightly decreased (Table 2, Movie S1 and S2). Moreover, morphological anomalies of the sperm flagella were easily observed, such as an unmasking, bending, or cracking axoneme, which recapitulated the flagellar phenotype of the infertile patient (Fig. 3, A and B). We further found that during the spermatogenic process, the abnormality of axoneme exposure had already occurred when the spermatozoa flagellum developed (fig. S5, A). Intriguingly, tail defects, including bending and cracking between the middle piece and the principal piece, were mainly present in the epididymal spermatozoa (fig. S5, B), suggesting that more severe flagellum breakage might occur during sperm movement.

The absence of Iqch impaired spermatogenesis in mice.

(A and B) The Papanicolaou staining (B) and SEM (C) results showing unmasking, bending, or cracking axoneme in the spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars in B, 5 μm; scale bars in C, 2.5 μm). The red arrowheads point to the axoneme abnormalities. (C) TEM revealing dilated intermembrane spaces of mitochondria and a normal annulus in the testicular and epididymal spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars, 500 nm). The red arrowheads point to the dilated intermembrane spaces of mitochondria. The green arrowheads point to the normal annulus. M, mitochondria.

We further performed TEM analysis to investigate the ultrastructural defects in the testicular and epididymal spermatozoa of the Iqch KO mice. In contrast to the WT mice, the spermatozoa of the Iqch KO mice showed dilated intermembrane spaces of mitochondria and even loss of some mitochondrial material (Fig. 3C). However, the annulus did not present significant differences between the Iqch KO and WT mice (Fig. 3C), indicating that the disruptive link between the flagellar middle piece and the principal piece might result from mitochondrial defects but not the annulus. The immunofluorescence staining of the mitochondrial marker (Solute carrier family 25 member 4: SLC25A4) and SEPT4 also revealed that the mitochondria of the spermatozoa of the Iqch KO mice were severely defective, and there was no significant difference in the annulus of the spermatozoa between the Iqch KO mice and WT mice, which was consistent with the proband (fig. S6, A and B).

Earlier version: Poor IVF outcomes in the Iqch KO male mice

Revised version: Poor in vitro fertilization (IVF) outcomes in the Iqch KO male mice

To evaluate the cause of the damaged fertility of the Iqch KO male mice, mature oocytes from WT female mice in their reproductive period were retrieved and placed into culture dishes with sperm from the Iqch KO mice and WT mice for IVF treatment. We found that 70% of the embryos had pronuclei in the WT group, while the fertilization rate of the Iqch KO group was dramatically reduced, accounting for 47% (Fig. 4A). Consequently, the rates of both two-cell embryos and blastocysts were significantly lower in the Iqch KO male mice than in the WT male mice (Fig. 4A). Noticeably, the inactive acrosome reaction was observed in most of the sperm from the Iqch KO mice (Fig. 4B), leading to the inability to cross the zona pellucida (fig. S7, A). Furthermore, abnormal acrosome development was evident in the testes of the Iqch KO mice using PNA staining (fig. S7, B). Similarly, the expression of PNA was abnormal in the mature sperm of the Iqch KO male mice (Fig. 4C). In the sperm and testis of the Iqch KO male mice and the proband, the PLCζ fluorescence signal was attenuated and abnormally localized (Fig. 4D, fig. S7, C and D). Taken together, these data suggested that Iqch might play an important role in the acrosomal formation, which is essential for fertilization.

Poor IVF treatment outcomes from using sperm from the Iqch KO mice.

Earlier version: (A) Representative pronucleus embryos, two-cell embryos, and blastocysts obtained from the WT mice and the Iqch KO mice. Significantly reduced fertilization rates, cleavage rates, and blastocyst formation rates were observed in the Iqch KO mice compared to the WT mice (n = 3 biologically independent WT mice and KO mice; (B) scale bars, 100 μm; Student’s t test; *p < 0.05; error bars, s.e.m). (B) The acrosome reaction rates in the capacitated spermatozoa from the WT mice and Iqch KO mice were determined by Coomassie brilliant blue staining. The acrosome reaction rates were reduced in the spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars, 5 μm; Student’s t-test; *p < 0.05; error bars, s.e.m). The green arrowheads indicate the reacted acrosomes. The red arrowheads indicate intact acrosomes. (C and D) PNA (C) and PLCζ (D) staining showing the abnormal acrosome morphology and aberrant PLCζ location and expression in the spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars, scale bars, 5 μm). The dotted box indicates the typical pattern of the PLCζ location and expression in the spermatozoa from the WT mice. PNA, peanut agglutinin; PLCζ, phospholipase C zeta 1.

Revised version: (A) Representative pronucleus embryos, two-cell embryos, and blastocysts obtained from the WT mice and the Iqch KO mice. Significantly reduced fertilization rates, cleavage rates, and blastocyst formation rates were observed in the Iqch KO mice compared to the WT mice (n = 3 biologically independent WT mice and KO mice; scale bars, 100 μm; Student’s t test; *p < 0.05; error bars, s.e.m). (B) The acrosome reaction rates in the capacitated spermatozoa from the WT mice and Iqch KO mice were determined by Coomassie brilliant blue staining. The acrosome reaction rates were reduced in the spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars, 5 μm; Student’s t-test; *p < 0.05; error bars, s.e.m). The green arrowheads indicate the reacted acrosomes. The red arrowheads indicate intact acrosomes. (C and D) PNA (C) and PLCζ (D) staining showing the abnormal acrosome morphology and aberrant PLCζ location and expression in the spermatozoa from the Iqch KO mice (n = 3 biologically independent WT mice and KO mice; scale bars, scale bars, 5 μm). The dotted box indicates the typical pattern of the PLCζ location and expression in the spermatozoa from the WT mice. PNA, peanut agglutinin; PLCζ, phospholipase C zeta 1.

RNA-binding proteins are the most relevant targets by which IQCH regulates spermatogenesis

Earlier version: To elucidate the molecular mechanism by which IQCH regulates male fertility, we performed LC-MS/MS analysis using mouse sperm lysates and detected 288 interactors of IQCH. Gene Ontology (GO) analysis of the IQCH-bound proteins revealed a particular enrichment in fertilization, sperm axoneme assembly, mitochondrial organization, calcium channel, and RNA processing (Fig. 5A). Intriguingly, 33 ribosomal proteins were identified (Fig. 5B), indicating that IQCH might be involved in protein synthesis. Using proteomic analysis on sperm from the Iqch KO mice, we further assessed key proteins that might be activated by IQCH. A total of 1,993 differential proteins were quantified, including 807 upregulated proteins and 1,186 downregulated proteins (Fig. 5C). GO analysis revealed that the significantly downregulated proteins were enriched in RNA processing, gene expression, mitochondrion biogenesis, and calcium ion regulation (Fig. 5D), which was consistent with the enrichment of the IQCH-bound proteins.

Revised version: To elucidate the molecular mechanism by which IQCH regulates male fertility, we performed liquid chromatography tandem mass spectrometry (LC-MS/MS) analysis using mouse sperm lysates and detected 288 interactors of IQCH (Data S1). Gene Ontology (GO) analysis of the IQCH-bound proteins revealed many particular enrichment pathways related to fertilization, sperm axoneme assembly, mitochondrion organization, calcium channel, and RNA processing (Fig. 5A). Intriguingly, 33 ribosomal proteins were identified (Fig. 5B), indicating that IQCH might be involved in protein synthesis. Using proteomic analysis on sperm from the Iqch KO mice, we further assessed key proteins that might be activated by IQCH. A total of 1,993 differential proteins were quantified, including 807 upregulated proteins and 1,186 downregulated proteins (Fig. 5C). GO analysis revealed that the significantly downregulated proteins were enriched in RNA processing, gene expression, mitochondrion biogenesis, and calcium ion regulation (Fig. 5D), which was consistent with the enrichment of the IQCH-bound proteins.

IQCH bound and activated male reproduction-related proteins in the sperm of the mice.

(A) Bubble plots of the GO analysis showing that the IQCH-interacting proteins are significantly enriched in spermatogenesis and RNA processing in three categories: biological process, cellular component, and molecular function. GO, gene ontology. (B) Thirty-three ribosomal proteins interacted with IQCH in the sperm from the mice. (C) A heatmap showing the differential protein results from the proteomic analysis of the sperm from the WT and Iqch KO mice. (D) Bubble plots showing the decreased enrichment of proteins related to the spermatogenetic process and RNA processing according to the GO analysis of the Iqch KO mice compared to the WT mice. (E) Venn diagram depicting the 76 overlapping proteins between the 1186 downregulated proteins in the sperm from the Iqch KO mice and the 288 proteins binding with IQCH. (F) Chord diagram showing 21 proteins involved in RNA binding, 8 proteins involved in mitochondrial function, and 4 proteins involved in calcium channel activity among the 76 overlapping proteins. HNRPAB showed the greatest reduction of expression in the Iqch KO mice.

Importantly, the cross-analysis revealed that 76 proteins were shared between the IQCH-bound proteins and the IQCH-activated proteins (Fig. 5E), implicating this subset of genes as direct targets. Among the 76 proteins, 21 were RNA-binding proteins (RBPs), 10 of which were suggested to be involved in spermatogenesis (Kuroda et al. 2000, Yang et al. 2012, Chapman et al. 2013, Fukuda et al. 2013, Legrand et al. 2019, Sechi et al. 2019, Liang et al. 2021, Tian and Petkov 2021, Wang et al. 2022, Xu et al. 2022); 8 were involved in mitochondrial function; and 4 were calcium channel activity-related proteins (Fig. 5F). We focused on SYNCRIP, HNRNPK, FUS, EWSR1, ANXA7, SLC25A4, and HNRPAB, the loss of which showed the greatest influence on the phenotype of the Iqch KO mice. Specifically, SYNCRIP, HNRNPK, FUS, EWSR1, and HNRPAB are RBPs that are linked to spermatogenesis by controlling mRNA translation or spermatid post-meiotic transcription (Kuroda et al. 2000, Fukuda et al. 2013, Sechi et al. 2019, Tian and Petkov 2021, Xu et al. 2022). ANXA7 is a calcium-dependent phospholipid-binding protein that is a negative regulator of mitochondrial apoptosis (Du et al. 2015). Loss of SLC25A4 results in mitochondrial energy metabolism defects in mice (Graham et al. 1997). We further confirmed the binding of IQCH between these proteins by the Co-IP assay (Fig. 6A) and substantiated their downregulation in the sperm of the Iqch KO mice by immunofluorescence staining and western blotting (Fig. 6B, fig. S8).

IQCH interacted with CaM to regulate the expression of HNRPAB and spermatogenesis.

Earlier version: (A) A Co-IP assay of the mouse sperm lysates revealed that the seven proteins most relevant to the phenotype of the Iqch KO mice bound with IQCH. (C) Western blotting showing the reduced expression of the seven proteins in the sperm from the Iqch KO mice compared to the WT mice. (C) qPCR analysis of RNA immunoprecipitation (RIP) using HNRPAB antibodies and IgG antibodies on the mouse sperm lysates showed that HNRPAB interacted with several RNAs associated with fertilization and axoneme assembly. The qPCR analysis of RIP performed on the sperm lysates from the Iqch KO mice revealed a decrease in the interaction between HNRPAB and the RNA targets compared to the WT mice (Student’s t-test; *p < 0.05; error bars, s.e.m). (D) The Co-IP assay showed the binding of IQCH and CaM in the WT sperm but not in the Iqch KO sperm. (E) The Co-IP assay showed that the decreased expression of HNRPAB was due to the reduced binding of IQCH and CaM by the knockout of IQCH or CaM. (F) The overexpression of IQCH and the simultaneous knockout of CaM or the overexpression of CaM and the simultaneous (D) knockout of IQCH in K562 cells further confirmed that the downregulation of HNRPAB was due to the diminished interaction between IQCH and CaM by western blotting analysis. (G) Analysis of the binding affinity between GFP-IQCH from the cell lysates (target) and recombinant CaM (ligand) by a microscale thermophoresis (MST) assay showing the interaction between IQCH and CaM. Their interaction was destroyed after the deletion of the IQ motif within IQCH. (H) The Co-IP assay of the HEK 293 cells co-transfected with WT-IQCH and WT-CaM plasmids or co-transfected with IQCH (△IQ) and WT-CaM plasmids showed that CaM interacted with the IQ motif. The downregulation of HNRPAB was due to the disrupted interaction between IQCH and CaM. Three independent experiments were performed.

Revised version: (A) A Co-IP assay of the mouse sperm lysates revealed that the seven proteins most relevant to the phenotype of the Iqch KO mice bound with IQCH. (B) Western blotting showing the reduced expression of the seven proteins in the sperm from the Iqch KO mice compared to the WT mice. (C) qPCR analysis of RNA immunoprecipitation (RIP) using HNRPAB antibodies and IgG antibodies on the mouse sperm lysates showed that HNRPAB interacted with several RNAs associated with fertilization and axoneme assembly. The qPCR analysis of RIP performed on the sperm lysates from the Iqch KO mice revealed a decrease in the interaction between HNRPAB and the RNA targets compared to the WT mice (Student’s t-test; *p < 0.05; error bars, s.e.m). (D) The Co-IP assay showed the binding of IQCH and CaM in the WT sperm. (E) The Co-IP assay showed that the decreased expression of HNRPAB was due to the reduced binding of IQCH and CaM by the knockout of IQCH or CaM. (F) The overexpression of IQCH and the simultaneous knockout of CaM or the overexpression of CaM and the simultaneous knockout of IQCH in K562 cells further confirmed that the downregulation of HNRPAB was due to the diminished interaction between IQCH and CaM by western blotting analysis. (G) Analysis of the binding affinity between GFP-IQCH from the cell lysates (target) and recombinant CaM (ligand) by a microscale thermophoresis (MST) assay showing the interaction between IQCH and CaM. Their interaction was destroyed after the deletion of the IQ motif within IQCH. (H) The Co-IP assay of the HEK 293 cells co-transfected with WT-IQCH and WT-CaM plasmids, co-transfected with IQCH (△IQ) and WT-CaM plasmids, or co-transfected with control and WT-CaM plasmids showed that CaM interacted with the IQ motif. The downregulation of HNRPAB was due to the disrupted interaction between IQCH and CaM. Three independent experiments were performed.

Earlier version: Among these interactors of IQCH, HNRPAB was the most significantly downregulated protein seen by proteomic analysis (Data S1), implying that HNRPAB might be the main target of IQCH. As an RBP, HNRPAB has been suggested to play an important role in spermatogenesis by regulating the translation of testicular mRNAs (Fukuda et al. 2013). We thus employed RNA-seq analysis using the sperm from the Iqch KO and WT mice to investigate the effects of RNA levels when IQCH is absent. Importantly, among the downregulated genes, most were related to male fertility, such as axoneme assembly, spermatid differentiation, flagellated sperm motility, and fertilization (fig. S9). We hypothesized that this downregulation is linked to HNRPAB binding. Considering the disrupted fertilization and axoneme assembly in the Igch KO mice, we chose the essential molecules involved in these two processes, including Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1, to confirm our speculation. We carried out RNA immunoprecipitation on formaldehyde cross-linked sperm followed by qPCR to evaluate the interactions between HNRPAB and Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1 in the KO and WT mice. As expected, the binding between HNRPAB and those important molecules was detected in the WT mice (Fig. 6C), supporting the functional role of HNRPAB in testicular mRNAs. Intriguingly, significantly decreased interactions between HNRPAB and those molecules were observed in the KO mice (Fig. 6C). Therefore, it is indicated that IQCH is involved in spermatogenesis by mainly regulating RBPs, especially HNRPAB, to further mediate essential mRNA expression in the testis.

Revised version: Among these interactors of IQCH, HNRPAB was the most significantly downregulated protein seen by proteomic analysis (Data S2), implying that HNRPAB might be the main target of IQCH. As an RBP, HNRPAB has been suggested to play an important role in spermatogenesis by regulating the translation of testicular mRNAs (Fukuda et al. 2013). We thus employed RNA-seq analysis using the sperm from the Iqch KO and WT mice to investigate the effects of RNA levels when IQCH is absent. Importantly, among the downregulated genes, most were related to male fertility, such as axoneme assembly, spermatid differentiation, flagellated sperm motility, and fertilization (fig. S9). We hypothesized that this downregulation is linked to HNRPAB binding. Considering the disrupted fertilization and axoneme assembly in the Igch KO mice, we chose the essential molecules involved in these two processes, including Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1, to confirm our speculation. We carried out RNA immunoprecipitation on formaldehyde cross-linked sperm followed by qPCR to evaluate the interactions between HNRPAB and Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1 in the KO and WT mice. As expected, the binding between HNRPAB and those important molecules was detected in the WT mice (Fig. 6C), supporting the functional role of HNRPAB in testicular mRNAs. Intriguingly, significantly decreased interactions between HNRPAB and those molecules were observed in the KO mice (Fig. 6C). Therefore, it is indicated that IQCH is involved in spermatogenesis by mainly regulating RBPs, especially HNRPAB, to further mediate essential mRNA expression in the testis.

The interaction of IQCH and CaM is a prerequisite for IQCH function

Earlier version: Given that IQCH is a calmodulin-binding protein, we hypothesized that IQCH regulates these key molecules by interacting with CaM (fig. S10, A and B). As expected, CaM interacted with IQCH, as indicated by LC-MS/MS analysis. We initially confirmed the binding of IQCH and CaM in the WT sperm but not in the KO sperm by the Co-IP assay (Fig. 6D). In addition, their colocalization was detected during spermatogenesis by immunofluorescence staining (fig. S10, C).

Revised version: Given that IQCH is a calmodulin-binding protein, we hypothesized that IQCH regulates these key molecules by interacting with CaM (fig. S10, A and B). As expected, CaM interacted with IQCH, as indicated by LC-MS/MS analysis. We initially confirmed the binding of IQCH and CaM in the WT sperm by the Co-IP assay (Fig. 6D). In addition, their colocalization was detected during spermatogenesis by immunofluorescence staining (fig. S10, C).

Earlier version: To confirm that the interaction of IQCH and CaM is the prerequisite for regulating HNRPAB expression, we downregulated IQCH or CaM in K562 cells, which express both IQCH and CaM. As expected, reduced binding of IQCH and CaM was observed in both cell lines with IQCH or CaM knocked down, with a concomitant decrease in the expression of HNRPAB (Fig. 6E, fig. S10, D). We then overexpressed IQCH and knocked down CaM or overexpressed CaM and knocked down IQCH in K562 cells to check the change in HNRPAB levels. Consistently, the two situations also showed a reduction in the expression of HNRPAB resulting from the diminished interaction between IQCH and CaM (Fig. 6F, fig. S10, E).

Revised version: We further confirmed the binding of IQCH and CaM in K562 cells, which express both IQCH and CaM (Fig. 6E). Moreover, a decrease in the expression of HNRPAB was observed when IQCH or CaM was knocked down (Fig. 6E, fig. S10D), suggesting that IQCH or CaM might regulate HNRPAB expression. We then overexpressed IQCH and knocked down CaM or overexpressed CaM and knocked down IQCH in K562 cells, and found that the reduced HNRPAB expression could not be rescued in both two situations (Fig. 6F, fig. S10, E), indicating that CaM or IQCH cannot mediate HNRPAB expression alone. These findings suggested that the HNRPAB expression might be regulated by the interaction of IQCH and CaM.

Earlier version: We further confirmed that the IQ motif of IQCH is required for CaM binding using a microscale thermophoresis (MST) assay. The binding affinity between IQCH (target) from cell lysates overexpressing GFP-IQCH plasmids and recombinant CaM (ligand) was enhanced with increasing concentrations of recombinant CaM (Fig. 6G). However, their binding was disrupted when the IQ motif of IQCH was deleted (Fig. 6G). The Co-IP assay verified the above findings (Fig. 6H). Not surprisingly, the expression of HNRPAB in cells co-transfected with IQCH (△IQ) and CaM plasmids was lower than that in the cells overexpressing the WT-IQCH and CaM plasmids (Fig. 6H). Collectively, our findings suggest that IQCH interacts with CaM via IQ-motif to manipulate the expression of important molecules, especially HNRPAB, to play a role in spermatogenesis (Fig. 7).

Revised version: We further confirmed that the IQ motif of IQCH is required for CaM binding using a microscale thermophoresis (MST) assay. The binding affinity between IQCH (target) from cell lysates overexpressing GFP-IQCH plasmids and recombinant CaM (ligand) was enhanced with increasing concentrations of recombinant CaM (Fig. 6G). However, their binding was disrupted when the IQ motif of IQCH was deleted (Fig. 6G). The Co-IP assay verified the above findings (Fig. 6H). Not surprisingly, the expression of HNRPAB in cells co-transfected with IQCH (△IQ) / CaM plasmids and control/ CaM plasmids was similar, but was lower than that in the cells overexpressing the WT-IQCH and CaM plasmids (Fig. 6H). Collectively, our findings suggest that IQCH interacts with CaM via IQ-motif to manipulate the expression of important molecules, especially HNRPAB, to play a role in spermatogenesis (Fig. 7).

Proposed model for the mechanisms of action underlying the involvement of IQCH in spermatogenesis.

IQCH interacts with CaM by IQ motif to regulate the expression of RNA binding proteins. RNA binding proteins, particularly HNRPAB, bind and regulate several RNAs that influence the development of acrosome, mitochondria, and axoneme, thereby playing a critical role in spermatogenesis.

Discussion

Throughout the evolution of divergent species, IQCH has been conserved, indicating its fundamental role in organisms. As a novel CaM-binding protein, IQCH was first identified in human and mouse testes and exclusively localized in spermatocytes and spermatids, suggesting its potential activity in spermatogenesis (Yin et al. 2005). However, since IQCH was identified, no other findings have supported its function in male reproduction. In this study, we found that the mutation of IQCH impaired male fertility, including flagellar morphological abnormalities and axoneme cracking. Aberrant ultrastructure in sperm with the IQCH mutation was associated with severely defective acrosomes and mitochondria. Furthermore, Iqch KO mice showed similar irregularities in the flagellum, especially the axoneme and mitochondria. We also identified that Iqch KO mice displayed different degrees of functional defects in acrosomes. Thus, our work supported the vital role of IQCH in flagellar and acrosome development.

Earlier version: To date, only four IQ motif-containing proteins, including IQCD, IQCF1, IQCG, and IQCN, have been suggested to participate in spermatogenesis (Harris et al. 2014, Fang et al. 2015, Zhang et al. 2019, Dai et al. 2022). Most IQ motif-containing proteins function in Ca2+dependent biological processes by binding to CaM (Chen et al. 2014, Harris et al. 2014, Fang et al. 2015); thus, IQCD, IQCF1, and IQCN are involved in fertilization and are relevant to the acrosome reaction, sperm capacitation, or manchette assembly (Fang et al. 2015, Zhang et al. 2019, Dai et al. 2022). In addition, because CaM has an impact on the actin cytoskeleton, Iqcg KO mice exhibit a detachment of sperm heads from tails (Harris et al. 2014). In our study, the Iqch KO mice also showed an impaired acrosome reaction, which caused reduced fertilization. Intriguingly, the axoneme breaks at the annulus and the mitochondrial defects were outstanding in the flagella of the Iqch KO mice, whose phenotypes were unexplored in the previously discovered IQ motif-containing proteins. Because CaM activation can stimulate actin cytoskeleton changes, it is reasonable that flagellum formation is defective when IQCH is absent. Ca2+ is a key player in the regulation of mitochondrial functions (Bravo-Sagua et al. 2017). Thus, the absence of IQCH leads to a failed interaction with CaM and disrupts normal Ca2+ signaling, which consequently causes mitochondrial defects in Iqch KO mice and patients with IQCH variants. Our findings suggest that the fertilization function is the main action of IQ motif-containing proteins, while the IQ motif-containing protein has its own role in spermatogenesis.

Revised version: To date, only four IQ motif-containing proteins, including IQCD, IQCF1, IQCG, and IQCN, have been suggested to participate in spermatogenesis (Harris et al. 2014, Fang et al. 2015, Zhang et al. 2019, Dai et al. 2022). Most IQ motif-containing proteins function in Ca2+dependent biological processes by binding to CaM (Chen et al. 2014, Harris et al. 2014, Fang et al. 2015); thus, IQCD, IQCF1, and IQCN are involved in fertilization and are relevant to the acrosome reaction, sperm capacitation, or manchette assembly (Fang et al. 2015, Zhang et al. 2019, Dai et al. 2022). In addition, because CaM has an impact on the actin cytoskeleton, Iqcg KO mice exhibit a detachment of sperm heads from tails (Harris et al. 2014). In our study, the Iqch KO mice also showed an impaired acrosome reaction, which caused reduced fertilization. Intriguingly, the axoneme breaks at the annulus and the mitochondrial defects were outstanding in the flagella of the Iqch KO mice, whose phenotypes were unexplored in the previously discovered IQ motif-containing proteins. Because CaM activation can stimulate actin cytoskeleton changes, it is reasonable that flagellum formation is defective when IQCH is absent. Ca2+ is a key player in the regulation of mitochondrial functions (Bravo-Sagua et al. 2017). Thus, the absence of IQCH leads to a failed interaction with CaM and disrupts normal Ca2+ signaling, which consequently causes mitochondrial defects in Iqch KO mice and patients with IQCH variants. Our findings suggest that the fertilization function is the main action of IQ motif-containing proteins, while each specific IQ motif-containing protein also has its own distinct role in spermatogenesis.

Earlier version: However, there are few publications regarding the underlying mechanism of IQ motif-containing proteins manipulating reproductive processes, except for a recent study that stated that IQCN interacted with CaM to regulate the expression of CaM-binding proteins, including manchette-related proteins (LZTFL1, KIF27, and RSPH6A), IFT family proteins and their motor proteins (IFT22, IFT43, IFT74, IFT81, IFT140, IFT172, WDR19, TTC21B, and DYNC2H1), and ribosomal protein family proteins (RPS5, RPS25, RPS27, and RPSA) (Dai et al. 2022). In our study, it was observed that IQCH regulates the expression of RBPs by their interactions. RBPs are a large class of proteins that assemble with RNAs to form ribonucleoproteins (RNPs). RBPs function at various stages of RNA processing, including alternative splicing, RNA modification, nuclear export, localization, stability, and translation efficiency (Corley et al. 2020). RNA processing in male reproductive biology is essential for the production of mature sperm, and many RBPs have been demonstrated to be indispensable during spermatogenesis (Morgan et al. 2021). Among the RBPs, HNRPAB, a heterogeneous nuclear RNP, showed the most significant reduction among the downregulated interactors of IQCH, suggesting that HNRPAB might be the most important employee of IQCH. It has been reported that HNRPAB plays a central role in spermatogenesis by regulating stage-specific translation of testicular mRNAs (Fukuda et al. 2013). As expected, a RIP assay using an anti-HNRPAB antibody revealed binding between HNRPAB and the mRNAs of several important genes (Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1) involved in spermatogenesis. We further found that the interaction of IQCH and CaM is a prerequisite for regulating HNRPAB expression. We thus hypothesized that IQCH regulates male fertility by binding to CaM and controls HNRPAB to manipulate the expression of key mRNAs involved in spermatogenesis.

Revised version: However, there are few publications regarding the underlying mechanism of IQ motif-containing proteins manipulating reproductive processes, except for a recent study that stated that IQCN interacted with CaM to regulate the expression of CaM-binding proteins, including manchette-related proteins (LZTFL1, KIF27, and RSPH6A), IFT family proteins and their motor proteins (IFT22, IFT43, IFT74, IFT81, IFT140, IFT172, WDR19, TTC21B, and DYNC2H1), and ribosomal protein family proteins (RPS5, RPS25, RPS27, and RPSA) (Dai et al. 2022). In our study, it was observed that IQCH regulates the expression of RBPs by their interactions. RBPs are a large class of proteins that assemble with RNAs to form ribonucleoproteins (RNPs). RBPs function at various stages of RNA processing, including alternative splicing, RNA modification, nuclear export, localization, stability, and translation efficiency (Corley et al. 2020). RNA processing in male reproductive biology is essential for theproduction of mature sperm, and many RBPs have been demonstrated to be indispensable during spermatogenesis (Morgan et al. 2021). Among the RBPs, HNRPAB, a heterogeneous nuclear RNP, showed the most significant reduction among the downregulated interactors of IQCH, suggesting that HNRPAB might be the most important downstream effector protein of IQCH. It has been reported that HNRPAB plays a central role in spermatogenesis by regulating stage-specific translation of testicular mRNAs (Fukuda et al. 2013). As expected, a RIP assay using an anti-HNRPAB antibody revealed binding between HNRPAB and the mRNAs of several important genes (Catsper1, Catsper2, Catsper3, Ccdc40, Ccdc39, Ccdc65, Dnah8, Irrc6, and Dnhd1) involved in spermatogenesis. We further found that the interaction of IQCH and CaM is a prerequisite for regulating HNRPAB expression. We thus hypothesized that IQCH regulates male fertility by binding to CaM and controls HNRPAB to manipulate the expression of key mRNAs involved in spermatogenesis.

In conclusion, our study identified a novel IQ motif-containing protein-coding gene, IQCH/Iqch, which is responsible for spermatogenesis in humans and mice, broadening the scope of known male infertility-associated genes thus further illustrating the genetic landscape of this disease. In addition, our study suggested an unexplored mechanism in which IQCH regulates the expression of key RBPs in spermatogenesis, especially HNRPAB, by interacting with CaM to play crucial roles in fertilization and axoneme assembly. We believe that our findings will provide genetic diagnostic markers and potential therapeutic targets for infertile males with IQCH pathogenic variants.

Materials and Methods

Study participants

A male infertility family was recruited from West China Second University Hospital, and healthy Chinese volunteers were enrolled as controls. The study was conducted according to the tenets of the Declaration of Helsinki, and ethical approval was obtained from the Ethical Review Board of West China Second University Hospital, Sichuan University. Each subject signed an informed consent form.

WES and Sanger sequencing

Peripheral blood samples were obtained from all subjects, and genomic DNA was extracted using a QIAamp DNA Blood Mini Kit (QIAGEN, Germany; 51126). The exomes of the subject were captured by Agilent SureSelect Human All Exon V6 Enrichment kits (Agilent, CA, USA) and then sequenced on a HiSeq X-TEN system (Illumina, CA, USA). All reads were mapped to the human reference sequence (UCSC Genome Browser hg19) using Burrows-Wheeler Alignment. After quality filtration with the Genome Analysis Toolkit, functional annotation was performed using ANNOVAR through a series of databases, including the 1000 Genomes Project, dbSNP, HGMD, and ExAC. PolyPhen-2, SIFT, MutationTaster, and CADD were used for functional prediction. Variant verification of IQCH in patients was confirmed by Sanger sequencing using the primers listed in table S1.

Minigene assay

The minigene splicing assay was used to explore the effect of the IQCH variant (c.387+1_387+10del) on splicing. WT-IQCH and Mut-IQCH sequences, including intron 3, exon 4, and intron 4, were PCR-amplified separately from the genomic DNA of the control and patient. The two amplified fragments were cloned into the minigene vector pSPL3 between the EcoRI and BamHI sites through the Basic Seamless Cloning and Assembly Kit (TransGen Biotech, China, CU201-02). After transfection into HEK293T cells by DNA and siRNA transfection reagent (Polypus, France, 101000046), the splicing patterns of the transcripts produced from the WT-IQCH and Mut-IQCH plasmids were analyzed by RT-PCR, gel electrophoresis, and Sanger sequencing. The primers used for the plasmid genomic amplification and the RT-PCR are listed in table S1.

Iqch knockout mice

The animal experiments were approved by the Experimental Animal Management and Ethics Committee of West China Second University Hospital, Sichuan University. All animal procedures complied with the Animal Care and Use Committee of Sichuan University. The mouse Iqch gene has 6 transcripts and is located on chromosome 9. According to the structure of the Iqch gene, exons 2∼3 of the Iqch-201 (ENSMUST00000042322.10) transcript were recommended as the knockout region. The region contains a 215 bp coding sequence, and deletion of this region was expected to disrupt IQCH protein function. The sgRNA sequences, synthesized by Sangon Biotech (Shanghai, China), are listed in table S2. The two complementary DNA oligos of each sgRNA target were annealed and ligated to the pUC57-sgRNA plasmid (Addgene, USA) for cloning. The recombinant plasmid was transformed into DH5α competent cells, and the positive clone was screened based on kanamycin resistance and sequencing. The recombinant plasmid was linearized and purified by phenol-chloroform extraction. Transcriptions of the sgRNAs in vitro were performed using the MEGAshortscript Kit (Ambion, USA, AM1354) and purified using the MEGAclear Kit (Ambion, USA, AM1908). Cas9 mRNA was purchased from TriLink BioTechnologies. The mouse zygotes were co-injected with the RNA mixture of Cas9 mRNA (∼50 ng/μl) and sgRNA (∼30 ng/μl). The injected zygotes were transferred into pseudopregnant recipients to obtain the F0 generation. DNA was extracted from the tail tissues of the 7-day-old offspring, and PCR amplification was carried out with genotyping primers. A stable F1 generation (heterozygous mice) was obtained by mating the positive F0 generation mice with wild-type C57BL/6JG-pt mice. Primers used for the genotyping are listed in table S2.

Mouse fertility testing

To confirm the fertility of the Iqch KO male mice, natural mating tests were conducted. Briefly, three Iqch KO and three littermate control sexually mature male mice (8-12 weeks old) were paired with two 6- to 8-week-old normal C57BL/6J females (each male was mated with two female mice) for 6 months. The vaginal plugs of the mice were examined every morning. Female mice with vaginal plugs were fed separately, and the number of pups per litter was recorded.

Added in the revised version: IVF

Removed in the revised version: In vitro fertilization

Earlier version: Eight-week-old C57BL/6J female mice were superovulated by injecting 5 IU of pregnant mare serum gonadotropin (PMSG), followed by 5 IU of human chorionic gonadotropin (hCG) 48 h later. Sperm was released from the cauda epididymis of 10-week-old male mice, and sperm capacitation was performed for 50 min using a TYH solution. The cumulus-oocyte complexes (COCs) were obtained from the ampulla of the uterine tube 14 h after the hCG injection. The ampulla was torn with a syringe needle, and the COCs were gently squeezed onto liquid drops of an HTF medium. The COCs were then incubated with ∼5 μl of the sperm suspension in HTF liquid drops at 37 °C under 5% CO2. After 6 h, the eggs were washed several times using an HTF medium to remove the cumulus cells and then transferred to liquid drops of a KSOM medium.

Revised version: Eight-week-old C57BL/6J female mice were superovulated by injecting 5 IU of pregnant mare serum gonadotropin (PMSG), followed by 5 IU of human chorionic gonadotropin (hCG) 48 h later. Sperm was released from the cauda epididymis of 10-week-old male mice, and sperm capacitation was performed for 50 min using a test-yolk buffer with HEPES (TYH) solution. The cumulus-oocyte complexes (COCs) were obtained from the ampulla of the uterine tube 14 h after the hCG injection. The ampulla was torn with a syringe needle, and the COCs were gently squeezed onto liquid drops of a human fallopian tube fluid (HTF) medium. The COCs were then incubated with ∼5 μl of the sperm suspension in HTF liquid drops at 37 °C under 5% CO2. After 6 h, the eggs were washed several times using an HTF medium to remove the cumulus cells and then transferred to liquid drops of a potassium simplex optimized medium(KSOM).

Acrosome reaction analysis

Earlier version: The spermatozoa were collected from the cauda epididymis and capacitated for 50 min in a TYH medium at 37 °C under 5% CO2. Highly motile sperm were collected from the upper portion of the medium, and 10 μM of calcium ionophore A23187 (Sigma-Aldrich, CA, USA, C7522) was added to induce the acrosome reaction. After 15 min, the spermatozoa were spotted on a glass microscope slide, dried, and fixed with 4% PFA for 10 min. The acrosomes were stained with Coomassie brilliant blue staining.

Revised version: The spermatozoa were collected from the cauda epididymis and capacitated for 50 min in a TYH medium at 37 °C under 5% CO2. Highly motile sperm were collected from the upper portion of the medium, and 10 μM of calcium ionophore A23187 (Sigma-Aldrich, CA, USA, C7522) was added to induce the acrosome reaction. After 15 min, the spermatozoa were spotted on a glass microscope slide, dried, and fixed with 4% paraformaldehyde (PFA) for 10 min. The acrosomes were stained with Coomassie brilliant blue staining.

Earlier version: Hematoxylin-eosin (H&E) staining

Revised version: H&E staining

For the staining of the testicular tissues, samples were dissected from the adult mice and fixed in 4% PFA overnight at 4 °C. The fixed tissues were embedded in paraffin, sectioned (5 μm thick), dewaxed, and rehydrated. The sections of the testis were stained with hematoxylin and an eosin solution (Beyotime, Shanghai, China, C0105M) before imaging using a microscope (Leica, Germany).

Immunofluorescence staining

For sperm immunostaining, fresh sperm samples were fixed with 4% PFA for 20 min at room temperature. For the staining of the testis tissues, heat-induced antigen retrieval was performed in a citrate antigen retrieval solution (Beyotime, Shanghai, China, P0081) according to the manufacturer’s requirements. After permeabilization with 1% Triton X-100 for 30 min, the sperm slides were blocked with 5% bovine serum albumin serum for 1 h at room temperature. Primary antibodies were added to the slide and incubated overnight at 4 °C. After washing three times with PBS, the slides were incubated with secondary antibodies for 1 h at room temperature. The nuclei were counterstained with DAPI dye (Sigma[Aldrich, CA, USA, 28718-90-3) and mounted with an antifade mounting medium. Images were captured with a laser-scanning confocal microscope (Olympus, Japan). Detailed information on the antibodies is provided in table S3.

qPCR

Earlier version: Total RNA was extracted from the mouse testes using an RNA Easy Fast Tissue/Cell Kit (Tiangen Biotech, China, 4992732). Approximately 0.3 mg of total RNA was converted into cDNA with the PrimeScript RT Reagent Kit (Takara, Japan, RR037A) according to the manufacturer’s instructions. The cDNAs were individually diluted 10-fold to be used as templates for the subsequent real-time fluorescence quantitative PCR (qPCR) with iTaq Universal SYBR Green Supermix (Bio-Rad Laboratories, USA, 1725124). Mouse Gapdh was used as an internal control. The mRNA expression of detected targets was quantified according to the 2-ΔΔCt method. The primers used for the qPCR are listed in table S1 and table S2.

Revised version: Total RNA was extracted from the mouse testes using an RNA Easy Fast Tissue/Cell Kit (Tiangen Biotech, China, 4992732). Approximately 0.3 mg of total RNA was converted into cDNA with the PrimeScript RT Reagent Kit (Takara, Japan, RR037A) according to the manufacturer’s instructions. The cDNAs were individually diluted 10-fold to be used as templates for the subsequent qPCR with iTaq Universal SYBR Green Supermix (Bio-Rad Laboratories, USA, 1725124). Mouse Gapdh was used as an internal control. The mRNA expression of detected targets was quantified according to the 2-ΔΔCt method. The primers used for the qPCR are listed in table S1 and table S2.

Papanicolaou staining

Semen samples were fixed in 4% paraformaldehyde, and the coated slides were air-dried before being rehydrated with 80%, 70%, and 50% ethanol and distilled water. The samples were then stained with Lea’s hematoxylin, rinsed with distilled water, and stained with G-6 orange stain and EA-50. Following the staining, the slides were dehydrated with ethanol and mounted.

Added in the revised version: SEM

Removed in the revised version: Scanning electron microscopy (SEM)

The spermatozoa were fixed in 2.5% phosphate-buffered glutaraldehyde (GA) (Zhongjingkeyi Technology, Beijing, China) at room temperature for 30 min and then deposited on coverslips. The coverslips were dehydrated via an ascending gradient of 50%, 70%, 95%, and 100% ethanol and air-dried. The specimens were then attached to specimen holders and coated with gold particles using an ion sputter coater before being viewed with a JSM-IT300 scanning electron microscope (JEOL, Tokyo, Japan).

Added in the revised version: TEM

Removed in the revised version: Transmission electron microscopy (TEM)

Precipitation of the spermatozoa was fixed with 2.5% (vol/vol) glutaraldehyde in a 0.1 M phosphate buffer (PB) (pH 7.4), washed two times in PB, and two times in ddH2O. Then, the tissues were immersed in 1% (wt/vol) OsO4 and 1.5% (wt/vol) potassium ferricyanide aqueous solution at 4 °C for 2 h. After washing, the samples were dehydrated through graded alcohol (30%, 50%, 70%, 80%, 90%, 100%, 10 min each) in pure acetone (10 min twice). The samples were then infiltrated in a graded mixture (3:1, 1:1, 1:3) of acetone and SPI-PON812 resin (21 ml SPO-PON812, 13 ml DDSA and 11 ml NMA), and then the pure resin was changed. The specimens were embedded in pure resin with 1.5% BDMA and polymerized for 12 h at 45 °C and 48 h at 60 °C. The ultrathin sections (70 nm thick) were sectioned with a microtome (Leica EM UC6, Germany), double-stained with uranyl acetate and lead citrate, and examined by a transmission electron microscope (TECNAI G2 F20, Philips, USA).

Proteomic quantification

Earlier version: Proteins were extracted from the sperm samples using a RIPA lysis buffer (Applygen, Beijing, China, C1053) containing 1 mM PMSF and protease inhibitors on ice. The supernatants were collected following centrifugation at 12,000 × g for 20 min. Protein concentrations were calculated by Bradford quantification and SDS-PAGE. An enzyme solution at a ratio of 1:20 of trypsin enzyme (μg) to substrate protein (μg) was added to 100 μg of the protein samples, vortexed, centrifuged at low speed for 1 minute, and incubated at 37 °C for 4 h. The peptide liquid obtained after salt removal was freeze-dried. The peptide sample was dissolved in 0.5 M TEAB and added to the corresponding iTRAQ labeling reagent and stored at room temperature for 2 h. The Shimadzu LC-20AB liquid phase system was used for purification, and the separation column was a 5 μm 4.6 × 250 mm Gemini C18 column for liquid phase separation of the sample. The dried peptide samples were reconstituted with mobile phase A (2% can, 0.1% FA) and centrifuged at 20,000 × g for 10 min, and the supernatant was taken for injection. The separation was performed by UltiMate 3000 UHPLC (Thermo Fisher). The sample was first enriched in a trap column and desalted and then entered into a self-packed C18 column. The peptides separated by the liquid phase chromatography were ionized by a nanoESI source and then passed to a tandem mass spectrometer Q-Exactive HF X (Thermo Fisher) for DDA (Data Dependent Acquisition) mode detection. The raw data were converted to mgf files for bioinformatics analysis, and protein identification from tandem mass spectra was performed by database searching (UniProt). The protein quantification process includes the following steps: protein identification, tag impurity correction, data normalization, missing value imputation, protein ratio calculation, statistical analysis, and results presentation. Proteins with a 1.5-fold change and a p-value (using Student’s t-test) less than 0.05 were defined as differentially expressed proteins.

Revised version: Proteins were extracted from the sperm samples using a radioimmunoprecipitation assay buffer (RIPA) lysis buffer (Applygen, Beijing, China, C1053) containing 1 mM phenylmethanesulfonyl fluoride (PMSF) and protease inhibitors on ice. The supernatants were collected following centrifugation at 12,000 × g for 20 min. Protein concentrations were calculated by Bradford quantification and sodium dodecyl sulfate polyacrylamide gel electrophoresis (SDS-PAGE). An enzyme solution at a ratio of 1:20 of trypsin enzyme (μg) to substrate protein (μg) was added to 100 μg of the protein samples, vortexed, centrifuged at low speed for 1 minute, and incubated at 37 °C for 4 h. The peptide liquid obtained after salt removal was freeze-dried. The peptide sample was dissolved in 0.5 M TEAB and added to the corresponding iTRAQ labeling reagent and stored at room temperature for 2 h. The Shimadzu LC-20AB liquid phase system was used for purification, and the separation column was a 5 μm 4.6 × 250 mm Gemini C18 column for liquid phase separation of the sample. The dried peptide samples were reconstituted with mobile phase A (2% can, 0.1% FA) and centrifuged at 20,000 × g for 10 min, and the supernatant was taken for injection. The separation was performed by UltiMate 3000 UHPLC (Thermo Fisher). The sample was first enriched in a trap column and desalted and then entered into a self-packed C18 column. The peptides separated by the liquid phase chromatography were ionized by a nanoESI source and then passed to a tandem mass spectrometer Q-Exactive HF X (Thermo Fisher) for DDA (Data Dependent Acquisition) mode detection. The raw data were converted to mgf files for bioinformatics analysis, and protein identification from tandem mass spectra was performed by database searching (UniProt). The protein quantification process includes the following steps: protein identification, tag impurity correction, data normalization, missing value imputation, protein ratio calculation, statistical analysis, and results presentation. Proteins with a 1.5-fold change and a p-value (using Student’s t-test) less than 0.05 were defined as differentially expressed proteins.

RNA sequencing

Earlier version: Total RNA (∼1 μg) was isolated from fresh spermatozoa from the WT mice (n=3) and the Iqch KO mice (n=3) using the RNAsimple Total RNA Kit (Tiangen Biotech, China, 4992858). A Ribo-Zero™ rRNA Removal Kit (MRZPL1224, llumina) was used to remove rRNA from the samples. RNA integrity was evaluated using the Agilent 2100 Bioanalyzer system (Agilent Technologies). An mRNA sequencing library for the RNA-seq was constructed using the TruSeq RNA Library Prep kit v2 (RS-122-2001, Illumina), followed by paired-end (2[×[100[bp) sequencing using the Illumina HiSeq 4000 Sequencing System (Illumina) at the Beijing Genomics Institute. Raw paired-end reads were filtered by FASTX-Toolkit. The quality of the reads was confirmed by FastQC (ver: 0.11.3). The raw data (raw reads) in the fastq format were processed through in-house Perl scripts. Feature Counts (ver: l.5.0-p3) were used to count the read numbers mapped to each gene. The FPKM of each gene was calculated based on the length of the gene and the read count mapped to the gene. Signaling pathway matching analysis for the differentially expressed gene (DEG) list was performed using KEGG Mapper in the Kyoto Encyclopedia of Genes and Genomes (KEGG, https://www.genome.jp/kegg/). The mapping of GO to DEGs was carried out using Blast2GO (ver: 4.1.9).

Revised version: Total RNA (∼1 μg) was isolated from fresh spermatozoa from the WT mice (n=3) and the Iqch KO mice (n=3) using the RNAsimple Total RNA Kit (Tiangen Biotech, China, 4992858). A Ribo-Zero™ rRNA Removal Kit (MRZPL1224, llumina) was used to remove rRNA from the samples. RNA integrity was evaluated using the Agilent 2100 Bioanalyzer system (Agilent Technologies). An mRNA sequencing library for the RNA-seq was constructed using the TruSeq RNA Library Prep kit v2 (RS-122-2001, Illumina), followed by paired-end (2[×[100[bp) sequencing using the Illumina HiSeq 4000 Sequencing System (Illumina) at the Beijing Genomics Institute. Raw paired-end reads were filtered by FASTX-Toolkit. The quality of the reads was confirmed by FastQC (ver: 0.11.3). The raw data (raw reads) in the fastq format were processed through in-house Perl scripts. Feature Counts (ver: l.5.0-p3) were used to count the read numbers mapped to each gene. The fragments per kilobase million (FPKM) of each gene was calculated based on the length of the gene and the read count mapped to the gene. Signaling pathway matching analysis for the differentially expressed gene (DEG) list was performed using Kyoto Encyclopedia of Genes and Genomes (KEGG) Mapper (https://www.genome.jp/kegg/). The mapping of GO to DEGs was carried out using Blast2GO (ver: 4.1.9).

Western blotting and Co-IP

Earlier version: Proteins were extracted from the cultured cells and mouse testes using a RIPA lysis buffer containing 1 mM PMSF and protease inhibitors (Applygen, Beijing, China, P1265) on ice. The supernatants were collected following centrifugation at 12,000 × g for 20 min. The proteins were electrophoresed in 10% SDS-PAGE gels and transferred to nitrocellulose membranes (GE Healthcare). The blots were blocked in 5% milk and incubated with primary antibodies overnight at 4 °C, followed by incubation with anti-rabbit or anti-mouse IgG H&L (HRP) (Abmart, Shanghai, China, M21002 and M21001) at a 1/10,000 dilution for 1 h. The signals were evaluated using a Super ECL Plus Western Blotting Substrate (Applygen, Beijing, China, P1050) and a Tanon-5200 Multi chemiluminescence imaging system (Tanon, Shanghai, China).

Revised version: Proteins were extracted from the cultured cells and mouse testes using a RIPA lysis buffer containing 1 mM PMSF and protease inhibitors (Applygen, Beijing, China, P1265) on ice. The supernatants were collected following centrifugation at 12,000 × g for 20 min. The proteins were electrophoresed in 10% SDS-PAGE gels and transferred to nitrocellulose membranes (GE Healthcare). The blots were blocked in 5% milk and incubated with primary antibodies overnight at 4 °C, followed by incubation with anti-rabbit or anti-mouse IgG heavy and light chain(H&L) (horseradish peroxidase, HRP) (Abmart, Shanghai, China, M21002 and M21001) at a 1/10,000 dilution for 1 h. The signals were evaluated using a Super ECL Plus Western Blotting Substrate (Applygen, Beijing, China, P1050) and a Tanon-5200 Multi chemiluminescence imaging system (Tanon, Shanghai, China).

For the Co-IP assays, the extracted proteins were incubated with primary antibodies overnight at 4 °C. The lysates were then incubated with 20 μl of Pierce™ Protein A/G-conjugated Agarose for 2 h at room temperature. The beads were washed with a washing buffer [50 mM Tris-HCl (pH 7.4), 0.1% Triton X-100, and 500 mM NaCl], eluted with a 1.2×SDS loading buffer, and boiled for 5 min at 95 °C. Finally, the products were separated using SDS-PAGE and analyzed using immunoblotting procedures. Detailed information on the antibodies used in the western blotting experiments is provided in table S3.

Antibody production

The peptide used for raising the anti-IQCH antibody was derived from amino acid residues 406–435 (KAEAATKIQATWKSYKARSSFISYRQKKWA) of mouse IQCH. The peptide coupled with keyhole limpet hemocyanin (KLH) (Sigma-Aldrich, USA, H7017) was dissolved in saline, emulsified with 1[ml of Freund’s complete adjuvant (Beyotime, China, P2036), and injected at multiple sites on the back of New Zealand white rabbits. The antiserum was collected within 2 weeks after the final injection.

Cell culture and transfection

Earlier version: We purchased the K562 cells (CRL-3344) and the HEK293T cells (CRL-11268) from the American Type Culture Collection. The K562 cells were cultured in Basic Roswell Park Memorial Institute 1640 Medium (Gibco, USA, C11875500BT) supplemented with 10% fetal bovine serum (Gibco, USA, 12483020). The HEK293T cells were cultured in Dulbecco’s Modified Eagle Medium (Gibco, USA, 11965092) supplemented with 10% fetal bovine serum (Gibco, USA, 12483020). The expression plasmids, including pcDNA3.1-Flag-CALM2, pCMV-MCS-3* flag-WT-IQCH, and pCMV-MCS-3* flag-IQCH(△IQ), and small interfering RNAs (siRNAs) of CALM2 and IQCH were constructed by Vigene Biosciences (Jinan, China). The target sequences of siRNA are listed in table S1. The plasmids and siRNA were transfected into cells with a jetPRIME transfection reagent (Polypus, France, 101000046) according to the manufacturer’s protocol.

Revised version: We purchased the K562 cells (CRL-3344) and the HEK293T cells (CRL-11268) from the American Type Culture Collection. The K562 cell line is a human leukemia cell line showing a high expression levels of IQCH and CaM, and is used for the knockdown experiments. The K562 cells were cultured in Basic Roswell Park Memorial Institute 1640 Medium (Gibco, USA, C11875500BT) supplemented with 10% fetal bovine serum (Gibco, USA, 12483020). The HEK293T cells were cultured in Dulbecco’s Modified Eagle Medium (Gibco, USA, 11965092) supplemented with 10% fetal bovine serum (Gibco, USA, 12483020). The expression plasmids, including pcDNA3.1-Flag-CALM2, pCMV-MCS-3* flag-WT-IQCH, and pCMV-MCS-3* flag-IQCH(△IQ), and small interfering RNAs (siRNAs) of CALM2 and IQCH were constructed by Vigene Biosciences (Jinan, China). The target sequences of siRNA are listed in table S1. The plasmids and siRNA were transfected into cells with a jetPRIME transfection reagent (Polypus, France, 101000046) according to the manufacturer’s protocol.

Earlier version: Microscale thermophoresis (MST) assay

Revised version: MST assay

The MST experiments were conducted on a Monolith NT.115 system (NanoTemper Technologies, Germany). Lysates of the HEK293T cells transfected with fluorescent GFP-IQCH or GFP-IQCH(△IQ) were normalized by raw fluorescence (count), diluted using an MST buffer (50 mM Tris-HCl pH 7.5, 150 mM NaCl, 10 mM MgCl2, 0.05% (v/v) Tween 20), and added to 16 PCR tubes (10 μl per tube). Then, the purified CaM was diluted using an MST buffer into 16 gradients. Ten microliters of different concentrations of CaM were mixed with 10 μl of fluorescent GFP-IQCH protein and GFP-IQCH(△IQ) and reacted in a dark box for 15 min at room temperature. The samples were added to monolith capillaries (NanoTemper, MO-L022) and subsequently subjected to MST analysis. The measurement protocol times were as follows: fluorescence before 5 sec, MST on 30 sec, fluorescence after 5 sec, delay 25 sec. The dissociation constant (Kd) was determined using a single-site model to fit the curve.

Statistical analysis

Data were compared for statistical significance using GraphPad Prism version 9.0.0 (GraphPad Software). The unpaired, two-tailed Student’s t-test was used for the statistical analyses. The data are presented as the mean ± SEM, and statistically significant differences are represented as *p<0.05.

Supporting information